READ ME:

Permissions for this test: 

- No talking until you have left the room.
- You may use the Internet during this exam.


To  begin the exam:

- **Rename the notebook to be of the form `<uni>_exam`. For example, mine would be `pl2648_exam`.**

When you are done with your exam:

1. Save this exam.
1. Download this exam from your Jupyter Notebook tab as an **`.ipynb`** file.
1. Upload/email/etc the **`.ipynb`** file to the submission platform designated by the exam proctor.

----

Please note, there are several cells in this Jupyter notebook that are empty and/or read only. Do not attempt to remove them or edit them. They are used in grading your notebook.

- DO remove the "Not Implemented" lines if you at all attempt the problem.
- DO test all cells to make sure they run in 30 seconds or less.

Finally, spelling counts! If you are told to create a function, class, variable by a specific name then do so with the name and capitalization exactly as given. Otherwise you will loose most points. Because of this, its best to copy and paste the names as they are given.

### Question: Basics

Build a function that can both filter an aggregate. By default, the function should just return the passed in list of items. If either a `filterer` or `aggregator` are passed in, they should be used. The `filterer` function should take an item and return true if the value should be included in the output list; ie. it is predicate function. The `aggregator` function should aggregate any items passed into the function and potentially filtered by the `filterer`.

<!-- grade_id: cell-3275940cea988440, points: 0, grade: false, solution: true -->
```python
def filter_and_agg(items: list, filterer=None, aggregator=None) -> list:
    '''
    Build a function that can both filter an aggregate. By default, the function should just return the passed in list of items. If either a `filterer` or `aggregator` are passed in, they should be used. The `filterer` function should take an item and return true if the value should be included in the output list; ie. it is predicate function. The `aggregator` function should aggregate any items passed into the function and potentially filtered by the `filterer`.
    '''
    if filterer is None and aggregator is None:
        return items
    elif filterer is None and aggregator is not None:
        result = aggregator(items)
        if result is not list:
            return [result]
        return aggregator(items)
    elif filterer is not None and aggregator is None:
        f_items = [] 
        for i in items:
            if filterer(i):
                f_items.append(i)
        return f_items
    else:
        f_items = []
        for i in items:
            if filterer(i):
                f_items.append(i)
        result = aggregator(f_items)
        if result is not list:
            return [result]
        return aggregator(result)
```

<!-- grade_id: cell-a20fc915c539a67f, points: 1, grade: true, solution: false -->
```python
### BEGIN TESTS
result = filter_and_agg(
    [1, 2, 3],
)
assert result == [1, 2, 3]
### END TESTS
```

<!-- grade_id: cell-29db57030a72d6b1, points: 2, grade: true, solution: false -->
```python
### BEGIN TESTS
result = filter_and_agg(
    [1, 2, 3],
    filterer=lambda x: x % 2,
)
print(result)
assert result == [1, 3]
### END TESTS
```

Output:
```
[1, 3]
```

<!-- grade_id: cell-666f3dce65b54cb9, points: 1, grade: true, solution: false -->
```python
### BEGIN TESTS
result = filter_and_agg(
    [1, 2, 3],
    aggregator=sum,
)
print(result)
assert result == [6]
### END TESTS
```

Output:
```
[6]
```

<!-- grade_id: cell-62b96da73d066686, points: 2, grade: true, solution: false -->
```python
### BEGIN TESTS
result = filter_and_agg(
    [1, 2, 3],
    filterer=lambda x: x % 2,
    aggregator=sum,
)
print(result)
assert result == [4]
### END TESTS
```

Output:
```
[4]
```

### Question: OOP

Develop a Queue. The class should have two methods besides the `__init__` method; `pop` and `push`. Item's pushed onto the queue should be removed from the queue via `pop`. Items should pass through the queue in a First In First Out pattern. A queue can be instantiated with a `max_size`. If a push occurs when the queue already contains `max_size` items, the `push` method should raise a `RuntimeError`. The contents of the error message can be anything.

<!-- grade_id: cell-6cc53c18dfc809ab, points: 0, grade: false, solution: true -->
```python
class Queue:
    def __init__(self, max_size: int=None):
        if max_size is not None and max_size < 1:
            self.q = []
        if max_size is None:
            self.q = []
            self.s =  float('inf')
        if max_size is not None and max_size > 0:
            self.s = max_size
            self.q = []
        #raise NotImplementedError()
    
    def push(self, value):
        if len(self.q) == self.s:
            raise RuntimeError("Should have raised RuntimeError")
        else:
            self.q.append(value)

        #raise NotImplementedError()
    
    def pop(self):
        # YOUR CODE HERE
        if len(self.q) == 0:
            return None
        else:
            return self.q.pop(0)
        #raise NotImplementedError()
```

<!-- grade_id: cell-f99e8f24067750fb, points: 1, grade: true, solution: false -->
```python
### BEGIN TESTS
q = Queue()
q.push(4)
assert q.pop() == 4
### END TESTS
```

<!-- grade_id: cell-71b6c3ff023d5de5, points: 1, grade: true, solution: false -->
```python
### BEGIN TESTS
q = Queue()
q.push(4)
q.push(8)
q.push(0)
assert q.pop() == 4
### END TESTS
```

<!-- grade_id: cell-72675c887573773f, points: 1, grade: true, solution: false -->
```python
### BEGIN TESTS
q = Queue()
q.push(4)
q.push(8)
q.push(0)
assert q.pop() == 4
assert q.pop() == 8
assert q.pop() == 0
### END TESTS
```

<!-- grade_id: cell-dab79708c71d45fe, points: 1, grade: true, solution: false -->
```python
### BEGIN TESTS
q = Queue(1)
q.push(4)
try:
    q.push(8)
    assert False, 'Should have raised RuntimeError'
except RuntimeError:
    pass
### END TESTS
```

<!-- grade_id: cell-d63b4d90f1a1dd32, points: 1, grade: true, solution: false -->
```python
### BEGIN TESTS
import random

q = Queue()
vals = []
for _ in range(10):
    i = random.randint(1, 100)
    q.push(i)
    vals.append(i)

for val in vals:
    assert val == q.pop()
### END TESTS
```

### Question: Regex

Construct a regex pattern that can parse simple math expressions. You can assume you only need to support `+`, `-`, `*`, `/` and integer operands. The function below should return the one-line pattern.

<!-- grade_id: cell-abd1593cac82c96e, points: 0, grade: false, solution: true -->
```python
def get_regex_pattern() -> str:
    return r'(\d+) *([\+\-\*\/]) *(\d+)'
    #raise NotImplementedError()
```

<!-- grade_id: cell-4080e33d2b132db1, points: 0, grade: true, solution: false -->
```python
### BEGIN TESTS
import re
### END TESTS
```

<!-- grade_id: cell-51d5fc75c4cd98d0, points: 1, grade: true, solution: false -->
```python
### BEGIN TESTS
result = re.match(get_regex_pattern(), '')
assert result is None, result
### END TESTS
```

<!-- grade_id: cell-aba7512346dde42b, points: 1, grade: true, solution: false -->
```python
### BEGIN TESTS
result = re.match(get_regex_pattern(), '1 + 1')
assert result.groups() == ('1', '+', '1'), result
### END TESTS
```

<!-- grade_id: cell-a5869387dfc8e6d0, points: 1, grade: true, solution: false -->
```python
### BEGIN TESTS
result = re.match(get_regex_pattern(), '1+1')
assert result.groups() == ('1', '+', '1'), result
### END TESTS
```

<!-- grade_id: cell-6253e03e88acf8ce, points: 1, grade: true, solution: false -->
```python
### BEGIN TESTS
result = re.match(get_regex_pattern(), '110 - 95')
assert result.groups() == ('110', '-', '95'), result
### END TESTS
```

<!-- grade_id: cell-218a08b6f2be92c8, points: 1, grade: true, solution: false -->
```python
### BEGIN TESTS
result = re.match(get_regex_pattern(), '23 * 77')
assert result.groups() == ('23', '*', '77'), result
### END TESTS
```

<!-- grade_id: cell-92fc5be6cbe7bcaa, points: 1, grade: true, solution: false -->
```python
### BEGIN TESTS
result = re.match(get_regex_pattern(), '23 / 99')
assert result.groups() == ('23', '/', '99'), result
### END TESTS
```

### Question: Bubble Sort

Write an implementation of bubble sort. Bubble sort compares to adjancent values and swaps them if they are out of order. The sort keeps passing through the data until no swaps have occurred during a pass.

Here is some psuedo-code to help the student:

```
While true:
    Step through the sequence.
        If two adjacent values are out of order, swap them.
    If there have been no swaps, stop looping.
```

<!-- grade_id: cell-d0c97b56e88a6448, points: 0, grade: false, solution: true -->
```python
def bubble_sort(seq: list):
    # YOUR CODE HERE
    for bl in range(len(seq)):
        #outer loop 
        for sl in range(0, len(seq) - bl - 1):
            #inner loop to compare and swap adjacent elements
            if seq[sl] > seq[sl + 1] :
                seq[sl], seq[sl + 1] = seq[sl + 1], seq[sl]
                #i and i + 1 are swapped
    #raise NotImplementedError()
```

<!-- grade_id: cell-7ff19d80cdb06bf1, points: 1, grade: true, solution: false -->
```python
### BEGIN TESTS
def get_result(list_: list):
    bubble_sort(list_)
    return list_
### END TESTS
```

<!-- grade_id: cell-d114ce6dac1c3484, points: 1, grade: true, solution: false -->
```python
### BEGIN TESTS
seq = []
result = get_result(seq)
assert result == [], result
### END TESTS
```

<!-- grade_id: cell-ad971a530598a253, points: 1, grade: true, solution: false -->
```python
### BEGIN TESTS
seq = [1, 2, 3]
result = get_result(seq)
assert result == [1, 2, 3], result
### END TESTS
```

<!-- grade_id: cell-796327d82b242e70, points: 1, grade: true, solution: false -->
```python
### BEGIN TESTS
seq = [3, 2, 1]
result = get_result(seq)
assert result == [1, 2, 3], result
### END TESTS
```

<!-- grade_id: cell-4cdcd85253e0f205, points: 1, grade: true, solution: false -->
```python
### BEGIN TESTS
import random

random.seed(0)

seq = random.choices(list(range(10)), k=10)
result = get_result(seq)
assert result == sorted(result), result
### END TESTS
```

```python

```

```python

```

```python

```