READ ME:

Permissions for this test: 

- No talking.
- This is a closed note, closed book, closed internet exam.
- You may have one jupyter notebook (this notebook) open for the duration of the exam.
- You may have one tab open solely for the act of submitting your exam.

To  begin the exam:

- **Rename the notebook to be of the form `<uni>_exam`. For example, mine would be `pl2648_exam`.**

When you are done with your exam:

1. Save this exam.
1. Download this exam as an **`.ipynb`** file.
1. Upload/email/etc the **`.ipynb`** file to the submission platform designated by the exam proctor.

----

Please note, there are several cells in this Jupyter notebook that are empty and read only. Do not attempt to remove them or edit them. They are used in grading your notebook.

- DO remove the "Not Implemented" lines if you at all attempt the problem
- DO test all cells to make sure they run in 30 seconds or less.

### Question

Create a `Building` class to represent a building.

The building should be instantiated with a given number of `windows`, `doors`, and `rooms` (given in that order) as positional arguments. A keyword argument named `fire_escapes` should also be accepted during instantiation and should be set to `None` by default. These given values should be stored on the object, under the same names, for retrieval after the object's construction.

For example:

```python
>>> b = Building(1, 2, 1)
>>> b.windows
1
```

The ratio of `windows` to `rooms` should be greater than 1.0. If that is not the case, a `ValueError` should be raised during instantiation. The error message should be `'Too few windows for given number of rooms.'`.

The ratio of `fire_escapes` to `rooms` should be greater than 0.5. If a value for `fire_escapes` is given, it should be validated according to this constraint. If no value for `fire_escapes` is given, one should be derived from the number of `rooms` given. The ratio of `fire_escapes` to `rooms` should be 0.5 in this case. Store this default ratio of 0.5 on the class as a class level attribute named `DEFAULT_FIRE_ESCAPE_RATIO` and use this attribute in any operations regarding `fire_escapes`. A `ValueError` should be raised during instantiation if the ratio is too low. The error message should be `'Too few fire escapes for given number of rooms.'`. 

All numbers (for `windows`, `doors`, `rooms`, and  `fire_escapes`) should be positive integers and should be validated as such. A `ValueError` should be raised if a non-conforming value is present at the end of instantiation. The message of the error should be of the form `'<attribute name> is <value> which is not an integer.'`

<!-- grade_id: oop-answer, points: 0, grade: false, solution: true -->
```python
import math


class Building:
    ### BEGIN SOLUTION
    
    DEFAULT_FIRE_ESCAPE_RATIO = 0.5

    def __init__(self, windows, doors, rooms, fire_escapes=None):
        self.windows = windows
        self.doors = doors
        self.rooms = rooms
        self.fire_escapes = fire_escapes
        
        self.derive_fire_escapes()
        
        self.validate()
    
    def derive_fire_escapes(self):
        if self.fire_escapes is not None:
            return
    
        self.fire_escapes = int(math.ceil(self.rooms * self.DEFAULT_FIRE_ESCAPE_RATIO))
    
    def validate(self):
        # Validate all parameters are integers
        attrs = ('windows', 'doors', 'rooms', 'fire_escapes')
        for attr in attrs:
            value = getattr(self, attr)
            if not isinstance(value, int):
                raise ValueError(f'{attr} is {value} which is not an integer.')
                
        # Validate window to room ratio
        if self.windows / self.rooms < 1:
            raise ValueError('Too few windows for given number of rooms.')
        
        # Validate fire escape ratio
        if self.fire_escapes / self.rooms < self.DEFAULT_FIRE_ESCAPE_RATIO:
            raise ValueError('Too few fire escapes for given number of rooms.')
    ### END SOLUTION
```

<!-- grade_id: oop-tests-init, points: 1, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
# Test init
assert Building(5, 2, 2).windows == 5
assert Building(5, 2, 2).doors == 2
assert Building(5, 2, 2).rooms == 2
assert Building(5, 2, 2).fire_escapes == 1
### END HIDDEN TESTS
```

<!-- grade_id: cell-1976862faf08794e, points: 1, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
# Test derive_fire_escapes
assert Building(5, 2, 4, 2).fire_escapes == 2
### END HIDDEN TESTS
```

<!-- grade_id: cell-7f28edc1af5caee6, points: 1, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
# Test int error messages
try:
    Building(4.5, 4, 4)
except ValueError as e:
    assert e.args[0] == 'windows is 4.5 which is not an integer.'
    
try:
    Building(5, 0.4, 4)
except ValueError as e:
    assert e.args[0] == 'doors is 0.4 which is not an integer.'
    
try:
    Building(5, 4, 4.2)
except ValueError as e:
    assert e.args[0] == 'rooms is 4.2 which is not an integer.'

try:
    Building(5, 4, 4, 2.5)
except ValueError as e:
    assert e.args[0] == 'fire_escapes is 2.5 which is not an integer.'
### END HIDDEN TESTS
```

<!-- grade_id: cell-44af3cce900e0c8e, points: 1, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
# Test window ratio
try:
    Building(1, 4, 4)
except ValueError as e:
    assert e.args[0] == 'Too few windows for given number of rooms.'
### END HIDDEN TESTS
```

<!-- grade_id: cell-bdd10244d13070f4, points: 1, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
# Test fire escape ratio
try:
    Building(5, 4, 4, 1)
except ValueError as e:
    assert e.args[0] == 'Too few fire escapes for given number of rooms.'
### END HIDDEN TESTS
```

<!-- grade_id: cell-5d16fcd763436720, points: 1, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
assert Building.DEFAULT_FIRE_ESCAPE_RATIO == 0.5
### END HIDDEN TESTS
```

### Question


Write a function called `common_letters` to find the common ASCII **letters** shared by two strings. Ignore case and your result should only contain lower case letters. You may only use one builtin function to perform this task but you may use it multiple times. You may use as many string methods as you like. You may perform as many operations between data structures as you like. You may perform an import outside of the function if necessary. The function should return an ordered list of letters. The body of the function can only be one line long.

For example:

```python
>>> common_letters('hi', 'hello')
['h']
```

<!-- grade_id: letters, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
import string

def common_letters(s1, s2):
    return sorted(set(s1.lower()) & set(s2.lower()) & set(string.ascii_lowercase))
### END SOLUTION
```

<!-- grade_id: cell-400e1f7d65286378, points: 1, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
common = common_letters("Hi â„ï¸! it's winter.", "Hello ðŸŒž, good to see you!")
assert common == ['e', 'h', 's', 't']
### END HIDDEN TESTS
```

<!-- grade_id: cell-931e198999083a99, points: 1, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
# Test only one line of code in function
assert common_letters.__code__.co_lnotab == b'\x00\x01'
### END HIDDEN TESTS
```

### Question

Create a generator function named `rev_fib` that, when called and cast to a list, returns the first N values in a reverse Fibonacci sequence.

For example: 

```python
>>> list(rev_fib(6))
... [0, -1, -1, -2, -3, -5]
```

<!-- grade_id: generator, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
def rev_fib(z):
    x = 0
    y = -1
    for _ in range(z):
        yield x
        y, x = y + x, y
### END SOLUTION
```

<!-- grade_id: gen-tests, points: 1, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
assert list(rev_fib(6)) == [0, -1, -1, -2, -3, -5]
assert list(rev_fib(100))[99] == -218922995834555169026
### END HIDDEN TESTS
```

<!-- grade_id: cell-6c182c209d2f71ab, points: 1, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
import inspect
assert inspect.isgeneratorfunction(rev_fib)
### END HIDDEN TESTS
```

### Question 

Write a function called `calibrate`. The calibration function should take an `OrderedDict`, a `predicate`, and an `update_func`. The `OrderedDict` passed to the function will have string keys and integer values. The function should update all values in the `OrderedDict` where the key for the value meets a predicate's constraints. Ie. `predicate(key) -> bool`. The update function should take a value (associated with a key that passed the predicate's constraints) and pass it through the `update_func` to arrive at the new value for the key. If the mapping passed to the `calibrate` function is not an `OrderedDict` a `TypeError` should be raised.

For example:

```python
>>> odict = collections.OrderedDict((
    ('a', 1),
    ('b', 2),
    ('c', 3),
))
>>> calibrate(odict, lambda x: x == 'b', lambda x: x **2)
>>> print(odict)
OrderedDict([('a', 1), ('b', 4), ('c', 3)])
```

<!-- grade_id: cell-9a0363302483e2d1, points: 0, grade: false, solution: true -->
```python
def calibrate(odict, predicate, update_func):
    ### BEGIN SOLUTION
    import collections
    if not isinstance(odict, collections.OrderedDict):
        raise TypeError
    
    for key in odict.keys():
        if predicate(key):
            odict[key] = update_func(odict[key])
    ### END SOLUTION
```

<!-- grade_id: cell-314d4f72425dc777, points: 2, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
import collections
odict = collections.OrderedDict((
    ('a', 1),
    ('b', 2),
    ('c', 3),
))

def predicate(key):
    return not (ord(key) - ord('a')) % 2

def update_func(value):
    return value ** value

assert calibrate(odict, predicate, update_func) == None, 'Function needs to mutate state'
assert odict == collections.OrderedDict([('a', 1), ('b', 2), ('c', 27)])
try:
    calibrate(dict(odict), predicate, update_func) == None
    assert False, 'Needs to raise error'
except TypeError:
    pass
except:
    assert False, 'Needs to be a TypeError'
### END HIDDEN TESTS
```

### Question

You've recieved a serialized JSON object from an API and have deserialized it using the standard library's `json` library. The object represents your geneology from a given ancestor downward. Assuming your name is Sally and your given ancestor is Janet, your geneology object would be as follows:

    geneology_object = {
        'husband': 'Craig', 
        'wife': 'Janet',
        'children': {
            'Chris': {
                'husband': 'Chris', 
                'wife': 'Jesse',
                'children': {
                    'Rebecca': {
                        'husband': 'Doug', 
                        'wife': 'Rebecca',
                    }
                }
            },
            'Wonda': {
                'husband': 'Kevin', 
                'wife': 'Wonda',
                'children': {
                    'Sally': {}
                }
            }
        }
    }


Write a function with the signature `get_generations_down(geneology_object, search_name, generations=0)` to recursively search for the number of generations between `search_name` and the eldest ancestor. If the name is not found, a `NameNotFoundError` should be raised by the recursive function.

Assuming the geneology object above, your function should behave as so:

    >>> get_generations_down(geneology_object, 'Chris')
    1
    >>> get_generations_down(geneology_object, 'Sally')
    2

<!-- grade_id: recursion, points: 0, grade: false, solution: true -->
```python
class NameNotFoundError(Exception):
    pass


def get_generations_down(geneology_object, search_name, generations=0):    
    ### BEGIN SOLUTION
    children = geneology_object.get('children', {})
    
    for child_name in children:
        if search_name == child_name:
            return generations + 1

        else:
            value = get_generations_down(children[child_name], search_name, (generations+1))
            if value:
                return value
    
    if generations == 0:
        raise NameNotFoundError('Name not found')
    ### END SOLUTION
```

<!-- grade_id: recursion-test-data, points: 0, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
geneology_object = {
    'husband': 'Craig', 
    'wife': 'Janet',
    'children': {
        'Chris': {
            'husband': 'Chris', 
            'wife': 'Jesse',
            'children': {
                'Rebecca': {
                    'husband': 'Doug', 
                    'wife': 'Rebecca',
                }
            }
        },
        'Wonda': {
            'husband': 'Kevin', 
            'wife': 'Wonda',
            'children': {}
        },
        'Mike': {
            'husband': 'Mike', 
            'wife': 'Rachael',
            'children': {
                'Eileen': {
                    'wife': 'Eileen',
                    'husband': 'Gary',
                    'children': {
                        'Patrick': {
                            'husband': 'Patrick',
                            'husband': 'Carl',
                        }
                    }
                } 
            }
        }
    }
}
### END HIDDEN TESTS
```

<!-- grade_id: cell-f56f214faab7f341, points: 2, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
assert get_generations_down(geneology_object, 'Patrick') == 3
assert get_generations_down(geneology_object, 'Eileen') == 2
### END HIDDEN TESTS
```

<!-- grade_id: cell-304f0d609b227e21, points: 1, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
try:
    get_generations_down(geneology_object, 'Paul')
    assert False, 'Should not have found Paul'
except NameNotFoundError: 
    pass
except:
    assert False, 'Should have raised NameNotFoundError'
### END HIDDEN TESTS
```

### Question


You work for the Department Of Housing in NYC. A new file has just been delivered to your desk via email. The file has row after row of cryptic housing information including an ID column. The ID column contains unique strings that are of the format `d[District ID (int)]b[Block ID (int)]l[Lot ID (int)]`. 

```
Example identifiers:
d17b4873l8390
d45b934l341
```

You need to parse these IDs for the ID parts and you decide that use of Python's regex library is your best tool. Complete the function below so that these IDs can be parsed by Python. Parse each sub ID (district, block, lot) into its own named group in the regex match object.

For example:

```
match = parse_id('d17b4873l8390')
match.group('district') ==> 17
match.group('block') ==> 4873
match.group('lot') ==> 8390
```

<!-- grade_id: re, points: 0, grade: false, solution: true -->
```python
import re

def parse_id(identifier):
    ### BEGIN SOLUTION
    pattern_string = r'd(?P<district>\d+)b(?P<block>\d+)l(?P<lot>\d+)'
    ### END SOLUTION
    return re.match(pattern_string, identifier)
```

<!-- grade_id: cell-2cf89dead6cbc6dd, points: 1, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
match = parse_id('d43b4637l27')
assert match is not None
### END HIDDEN TESTS
```

<!-- grade_id: cell-3a1cb82b63055930, points: 1, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
assert match.group('district') == '43'
assert match.group('block') == '4637'
assert match.group('lot') == '27'
### END HIDDEN TESTS
```